<?xml version="1.0" encoding="UTF-8"?>
<quiz>
    
        <!-- category: 1 -->
        <question type= "category" >
            <category>
                <text> TestCat1</text>
            </category>
            <info format="html">
                <text></text>
            </info>
            <idnumber></idnumber>
        </question>
    
        <!-- category: 2 -->
        <question type= "category" >
            <category>
                <text> TestCat1\TestSubCat1</text>
            </category>
            <info format="html">
                <text></text>
            </info>
            <idnumber></idnumber>
        </question>
    
        <!-- category: 3 -->
        <question type= "category" >
            <category>
                <text> TestCat2</text>
            </category>
            <info format="html">
                <text></text>
            </info>
            <idnumber></idnumber>
        </question>
    
        <!-- category: 4 -->
        <question type= "category" >
            <category>
                <text> TestCat2\TestSubCat2</text>
            </category>
            <info format="html">
                <text></text>
            </info>
            <idnumber></idnumber>
        </question>
    

    
        <question type= "category" >
            <category>
                <text> TestCat1</text>
            </category>
            <info format="html">
                <text></text>
            </info>
            <idnumber></idnumber>
        </question>
    
        <question type= "category" >
            <category>
                <text> TestCat1\TestSubCat1</text>
            </category>
            <info format="html">
                <text></text>
            </info>
            <idnumber></idnumber>
        </question>
    
        <question type= "category" >
            <category>
                <text> TestCat2</text>
            </category>
            <info format="html">
                <text></text>
            </info>
            <idnumber></idnumber>
        </question>
    
        <question type= "category" >
            <category>
                <text> TestCat2\TestSubCat2</text>
            </category>
            <info format="html">
                <text></text>
            </info>
            <idnumber></idnumber>
        </question>
    



    
        <!-- question: 1  -->
        <question type="coderunner">
        <name>
            <text>A rectangle's a</text>
        </name>
        <questiontext format="html">
          <text>A rectangle's area can be calculated using the following formula:

$$area = base * height$$

Write a program that prompts the user to enter the base and height of a rectangle in centimetres.  The program should
then print the area of the rectangle to two decimal places.

Refer to the example testcases below for the required prompts and format of the message to be displayed with rectangle area.

Note: The two values entered can be floating point values and will be valid non-negative numbers.</text>
        </questiontext>
        <generalfeedback format="html">
          <text></text>
        </generalfeedback>
        <defaultgrade>15</defaultgrade>
        <coderunnertype>python3_w_input</coderunnertype>
        <template>
<![CDATA[
__saved_input__ = input
def input(prompt=''):
    s = __saved_input__(prompt)
    print(s)
    return s

{{ STUDENT_ANSWER }}


SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]>
        </template>
        <answer>base = float(input("Enter the base of the rectangle in centimetres: "))
height = float(input("Enter the height of the rectangle in centimetres: "))

area = base * height

print(f"The area of the rectangle is {area:.2f}cm^2")</answer>
        <testcases>
        
            <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
            <testcode>
                    <text></text>
            </testcode>
            <stdin>
                    <text>12.5
5.5</text>
            </stdin>
            <expected>
                    <text>Enter the base of the rectangle in centimetres: 12.5
Enter the height of the rectangle in centimetres: 5.5
The area of the rectangle is 68.75cm^2</text>
            </expected>
            <extra>
                    <text></text>
            </extra>
                <display>
                    <text>SHOW</text>
                </display>
            </testcase>
        
            <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
            <testcode>
                    <text></text>
            </testcode>
            <stdin>
                    <text>100
200</text>
            </stdin>
            <expected>
                    <text>Enter the base of the rectangle in centimetres: 100
Enter the height of the rectangle in centimetres: 200
The area of the rectangle is 20000.00cm^2</text>
            </expected>
            <extra>
                    <text></text>
            </extra>
                <display>
                    <text>SHOW</text>
                </display>
            </testcase>
        
            <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
            <testcode>
                    <text></text>
            </testcode>
            <stdin>
                    <text>0
10</text>
            </stdin>
            <expected>
                    <text>Enter the base of the rectangle in centimetres: 0
Enter the height of the rectangle in centimetres: 10
The area of the rectangle is 0.00cm^2</text>
            </expected>
            <extra>
                    <text></text>
            </extra>
                <display>
                    <text>SHOW</text>
                </display>
            </testcase>
        
        </testcases>
      </question>
    
</quiz>